---
title       : OLS Basics 
description : This section contains exercises dealing with the simple linear regression model and estimation using ordinary least squares. 
attachments :
  slides_link : https://github.com/Emwikts1970/URFITE_DC/raw/master/Econometrics

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:35167f113b
## Multiple Regression: Boston Housing Data



*** =instructions
- Attach the package `MASS`
- Load the Boston housing data set (`Boston`)
- Regress the median housing value in a destrict (`medv`) on the average age of the buildings (`age`) and the crime rate (`crim`) 


*** =hint
Have a look at the plot. What can you say about the dispersion of observations?

*** =pre_exercise_code
```{r}
```

*** =sct
```{r}
```
